# NOTE importance of itertools

 `itertools.product()`

```python
import itertools
from itertools import product
```

# Sets
Read the notebook on sets before attempting these exercises

# Problem 1

De Morgan's first law states the following for any two sets $A$ and $B$
$$(A\cup B)^c = A^c\cap B^c$$

In the following two exercises we calculate $(A\cup B)^c$ in two different ways. Both functions must take $A$, $B$ and the universal set $U$ as their inputs.

## Exercise 1.1


Write the function **complement_of_union** that first determines $A\cup B$ and then evaluates the complement of this set. Output the tuple: $\begin{pmatrix}A\cup B,\, (A\cup B)^c\end{pmatrix}$.



<font  style="color:blue"> **Code**</font>
```python
A = {1, 2, 3}
B = {3, -6, 2, 0}
U = {-10, -9, -8, -7, -6, 0, 1, 2, 3, 4}
complement_of_union(A, B, U)
```

<font  style="color:magenta"> **Output**</font>
```
({-6, 0, 1, 2, 3}, {-10, -9, -8, -7, 4})
```


```python
# modify this cell

def complement_of_union(A, B, U):
    # inputs: A, B and U are of type 'set'
    # output: a tuple of the type (set, set)
    
    ### BEGIN SOLUTION
    C = A|B
    return (C, U-C)   
    ### END SOLUTION
```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
U = A|B|{-3, 7, 10, -4}
assert( complement_of_union(A, B, U) == ({-6, 0, 1, 2, 3, 4, 5, 8, 9}, {-4, -3, 7, 10})  )

### BEGIN HIDDEN TESTS
A = { -6, 2, 3, 4, 5, 9}
B = {0, -6, 5, 9, 11, 13}
U = A|B|{-3, 7, 12, -2, -4}
assert( complement_of_union(A, B, U) == ({-6, 0, 2, 3, 4, 5, 9, 11, 13}, {-4, -3, -2, 7, 12})   )
### END HIDDEN TESTS
```

## Exercsise 1.2

Write the function **intersection_of_complements** that first determines $A^c$ and $B^c$ and then evaluates the intersection of their complements. Output the tuple: $\begin{pmatrix}A^c, \,  A^c\cap B^c\end{pmatrix}$

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2, 3}
B = {3, -6, 2, 0}
U = {-10, -9, -8, -7, -6, 0, 1, 2, 3, 4}
intersection_of_complements(A, B, U)
```

<font  style="color:magenta"> **Output**</font>
```
({-10, -9, -8, -7, -6, 0, 4}, {-10, -9, -8, -7, 4})
```


```python
# modify this cell

def intersection_of_complements(A, B, U):
    # inputs: A, B and U are of type 'set'
    # output: a tuple of the form (set, set)
    
    ### BEGIN SOLUTION
    A_c = U - A
    B_c = U - B
    return (A_c, A_c & B_c)   
    ### END SOLUTION
```

<!-- grade_id: ex2, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
U = A|B|{-3, 7, 10, -4}
assert(  intersection_of_complements(A, B, U) == ({-6, -4, -3, 0, 7, 8, 9, 10}, {-4, -3, 7, 10})  )

### BEGIN HIDDEN TESTS
A = { -6, 2, 3, 4, 5, 9}
B = {0, -6, 5, 9, 11, 13}
U = A|B|{-3, 7, 12, -2, -4}
assert(  intersection_of_complements(A, B, U) == ({-4, -3, -2, 0, 7, 11, 12, 13}, {-4, -3, -2, 7, 12})  )
### END HIDDEN TESTS
```

# Problem 2

This problem illustrates a property of cartesian products of unions of two or more sets. For four sets $A$, $B$, $S$ and $T$, the following holds:

$$(A\cup B)\times(S\cup T) = (A\times S)\cup(A\times T)\cup(B\times S)\cup(B\times T)$$

Write the following functions to determine $(A\cup B)\times(S\cup T)$ in two different ways.


## Exercies 2.1

Write function **product_of_unions** that first determines $(A\cup B)$ and $(S\cup T)$ and then evaluates the cartesian products of these unions. Output the tuple $\begin{pmatrix}(A\cup B),\,  (A\cup B)\times(S\cup T)\end{pmatrix}$.

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2}
B = {1, 3}
S = {-1, 0}
T = {0, 10}
product_of_unions(A, B, S, T)
```


<font  style="color:magenta"> **Output**</font>
```
({1, 2, 3},
 {(1, -1),
  (1, 0),
  (1, 10),
  (2, -1),
  (2, 0),
  (2, 10),
  (3, -1),
  (3, 0),
  (3, 10)})
```


```python
# modify this cell

def product_of_unions(A, B, S, T):
    # inputs: A, B, S and T are sets
    # output: a tuple of the type (set, set)
    ### BEGIN SOLUTION
    from itertools import product
    R = set()
    for x in product(A|B, S|T):
        R.add(x)
    return (A|B, R) 
    ### END SOLUTION

# modify this cell
```

<!-- grade_id: ex3, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {5}
B = {5, 6}
S = {-1, 0, 1}
T = {1, 2}
assert( product_of_unions(A, B, S, T) == \
       ({5, 6}, {(5, -1), (5, 0), (5, 1), (5, 2), (6, -1), (6, 0), (6, 1), (6, 2)})   )

### BEGIN HIDDEN TESTS
A = {0, 5}
B = {5, 6}
S = {-1, 0, 1, 2}
T = {1, 6}
assert( product_of_unions(A, B, S, T) == \
 (set([0, 5, 6]), set([(0, 1), (5, -1), (0, 0), (6, 6), (5, 6), (6, 1), (6, 0), (6, -1), (0, 6), (6, 2), (5, 0), (5, 1), (0, -1), (5, 2), (0, 2)]))  )   
### END HIDDEN TESTS
```

## Exercise 2.2

Write a function **union_of_products** that first determines $(A\times S)$ and the other three cartesian products that appear on the right hand side of the identity above, then evaluates the union of these cartesian products. Output the tuple $\begin{pmatrix}(A\times S),\,  (A\times S)\cup(A\times T)\cup(B\times S)\cup(B\times T)\end{pmatrix}$.

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2}
B = {1, 3}
S = {-1, 0}
T = {0, 10}
union_of_products(A, B, S, T)
```


<font  style="color:magenta"> **Output**</font>
```
({(1, -1), (1, 0), (2, -1), (2, 0)},
 {(1, -1),
  (1, 0),
  (1, 10),
  (2, -1),
  (2, 0),
  (2, 10),
  (3, -1),
  (3, 0),
  (3, 10)})
```

```python
# modify this cell

def union_of_products(A, B, S, T):
    # inputs: A, B, S and T are sets
    # output: a tuple of the type (set, set)
    ### BEGIN SOLUTION
    from itertools import product
    P1 = set()
    P2 = set()
    P3 = set()
    P4 = set()
    for x in product(A, S):
        P1.add(x)
    for x in product(A, T):
        P2.add(x)
    for x in product(B, S):
        P3.add(x)
    for x in product(B, T):
        P4.add(x)
    return (P1, P1|P2|P3|P4)
    ### END SOLUTION

```

<!-- grade_id: ex4, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {5}
B = {5, 6}
S = {-1, 0, 1}
T = {1, 2}
assert( union_of_products(A, B, S, T) == \
        ({(5, -1), (5, 0), (5, 1)}, \
         {(5, -1), (5, 0), (5, 1), (5, 2), (6, -1), (6, 0), (6, 1), (6, 2)})  \
      )

### BEGIN HIDDEN TESTS
A = {0, 5}
B = {5, 6}
S = {-1, 0, 1, 2}
T = {1, 6}
assert( union_of_products(A, B, S, T) == \
(set([(0, 1), (5, -1), (0, 0), (5, 0), (5, 1), (0, -1), (5, 2), (0, 2)]), set([(0, 1), (5, -1), (0, 0), (6, 6), (5, 6), (6, 1), (6, 0), (6, -1), (0, 6), (6, 2), (5, 0), (0, -1), (5, 1), (5, 2), (0, 2)])) )
### END HIDDEN TESTS
```

```python







```